**Question 1** (this is Problem 1.3 from Gregory) 

We saw in class (and in Section 1.4.1) that based
on the saliva test result and the prior information, the probability that the person had the
unknown disease (UD) was found to be 0.43%. Subsequently, the same person received an
independent blood test for UD and again tested positive. If the false-negative rate for this
test is 1.4% and the false-positive rate is 0.5%, what is the new probability that the person
has UD on the basis of both tests?

<!-- grade_id: cell-23dd384285143cb6, points: 3, grade: true, solution: true -->
You can write your solutions down either 1) here in Jupyter using markdown
and Latex or 2) handwritten on paper and scanned/submitted to Canvas.

YOUR ANSWER HERE

**Question 2**

(a) Write a function that simulates a sequence of $n$ coin flips with probability of
heads $H$.  i.e.,

<code>flips = flipOneSet(n, H)</code>

where <code>flips</code> is an array of ”H” and ”T” for heads and tails, respectively

(b) Write a function that takes the array from above and returns the number of
heads and tails. i.e.,

<code>nheads, ntails = countFlips(flips)</code>

(c) Write a function that takes the array from part (a) and counts the number of
occurrences of $X$ consecutive heads. 
For example, in a sequence of 10 flips for which the
result is \[H, H, H, T, H, T, H, H, H, T\], the result for $X = 3$ should be 2; there are 3
instances of 2 consecutive heads (1-2-3 and 7-8-9).  The result for $X = 2$ should also be
2 (rather than 4). i.e.,

<code>nconsecutive = countConsecutives(flips, X)</code>

(d) Using the functions above, write another function that calculates the odds for a
sequence of $n$ coin flips with $H$ to result in at least one instance of $X$ consecutive heads.

<code>prob = probConsecutives(n, H, X)</code>

(HINT: You can do this by running flipOneSet multiple times (they need to be ”sufficiently” large) and measuring the fraction of occurrences.

<!-- grade_id: cell-21b8f6813cee61ba, points: 0, grade: false, solution: true -->
```python
# Question 2a
import numpy as np

def flipOneSet(n, H):
    """Return a numpy array of length n simulating a series of coin flips.
    
    Each item in the array is either 'H' or 'T'.
    The probability for each item to be 'H' is H.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-799a6a5fed9e5988, points: 2, grade: true, solution: false -->
```python
print("10 flips, fair coin: ", flipOneSet(10, 0.5))
print("Another 10 flips, fair coin:", flipOneSet(10, 0.5))
print("10 flips, two-sided heads: ", flipOneSet(10, 1.0))
print("15 flips, unfair coin: ", flipOneSet(15, 0.8))
print("1 flip, unfair coin: ", flipOneSet(1, 0.8))

# Run a few sanity checks to make sure the outputs are sensible.

flips = flipOneSet(100, 0.5)
assert isinstance(flips, np.ndarray)  # Don't return a python list.
assert len(flips) == 100              # Check that length of array is right.
assert not np.all(flips == 'H')       # Chances are extremely small that 100 flips are all H.
assert not np.all(flips == 'T')       # ... or T
flips2 = flipOneSet(100, 0.5)
assert not np.all(flips2 == flips)    # Doing it again gives a different result.

# Note: there are more hidden tests that are used for grading!

```

<!-- grade_id: cell-cf2a58e2162ff112, points: 0, grade: false, solution: true -->
```python
# Question 2b

def countFlips(flips):
    """Given an array of flips (all of which are either 'H' or 'T'), return the number of heads and tails.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4eed40f6866dc45e, points: 2, grade: true, solution: false -->
```python
flips = flipOneSet(15, 0.8)
nheads, ntails = countFlips(flips)
print(f"15 flips, unfair coin has {nheads} H and {ntails} T")

flips = flipOneSet(1000, 0.8)
nheads, ntails = countFlips(flips)
print(f"1000 flips, unfair coin has {nheads} H and {ntails} T")

assert nheads + ntails == 1000  # Every flip is either H or T

# For 1000 flips with p_H=0.8, nheads should be pretty close to 0.8 x 1000 = 800.
assert 700 < nheads < 900  # Chance of this failing should be extremely small.

```

<!-- grade_id: cell-3b9b40a520e77599, points: 0, grade: false, solution: true -->
```python
# Question 2c

def countConsecutives(flips, X):
    """Given an array of flips (all of which are either 'H' or 'T'), return how many times a run
    of X consecutive heads occurs.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e4cb2b4e28ca63c0, points: 3, grade: true, solution: false -->
```python
flips = np.array(['H', 'H', 'H', 'T', 'H', 'T', 'H', 'H', 'H', 'T'])
n_HHH = countConsecutives(flips, 3)
print(f"There are {n_HHH} occurrences of HHH in {flips}")

assert n_HHH == 2

# There is a potential ambiguity about how many HH appear in this sequence.  
# It could be 2 or 4, depending on whether you want HHH to count as having 1 or 2
# occurrences of HH.  I.e. whether overlapping runs count separately or not.
# The question says to not count overlapping runs, but we actually accept either.

n_HH = countConsecutives(flips, 2)
print(f"There are {n_HH} occurrences of HH in {flips}")
assert n_HH == 2 or n_HH == 4

n_HHHH = countConsecutives(flips, 4)
print(f"There are {n_HHHH} occurrences of HHHH in {flips}")
assert n_HHHH == 0

# Make sure it works with small arrays
assert countConsecutives(np.array(['H']), 1) == 1
assert countConsecutives(np.array(['H']), 2) == 0

```

<!-- grade_id: cell-192c408fbb2649b4, points: 0, grade: false, solution: true -->
```python
# Question 2d

def probConsecutives(n, H, X):
    """Estimate the probability of getting a run of X consecutive heads in a sequence of 
    n coin tosses where each toss has a probability H of landing heads.
    """
    # Hint: run many simulations of doing n and keep track of how often there is at
    #       least one occurrence of getting X consecutive heads.

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-756ed36d8b0abe8f, points: 3, grade: true, solution: false -->
```python
p1 = probConsecutives(10, 0.5, 3)
print(f'The chance of getting 3 consecutive heads in 10 flips of a fair coin = {p1}')

p2 = probConsecutives(10, 0.5, 4)
print(f'The chance of getting 4 consecutive heads in 10 flips of a fair coin = {p2}')

# Check a few where we can know the expected answer
p3 = probConsecutives(3, 0.5, 3)  # Should be close to 0.5^3 = 0.125
print(f'The chance of getting 3 consecutive heads in 3 flips of a fair coin = {p3}')
assert np.isclose(p3, 0.125, atol=0.02)

p4 = probConsecutives(5, 1.0, 4)  # Should be exactly 1
print(f'The chance of getting 4 consecutive heads in 5 flips of a two-headed coin = {p4}')
assert p4 == 1

p5 = probConsecutives(10, 0.0, 2)  # Should be exactly 0
print(f'The chance of getting 2 consecutive heads in 10 flips of a two-tailed coin = {p5}')
assert p5 == 0

```

**Question 3**

We saw in class that the posterior probability distribution
function (PDF) for $H$ = probablity of heads for a sequence of three coin flips $D = h, h, t$ is
proportional to $H^2(1 − H)$. i.e.,
$$
p(H|D, I) = A H^2(1 − H).
$$

(a) Calculate the constant $A$ that normalizes the PDF, i.e.,
$$
\int_0^1 p(H|D, I) dH = 1.
$$

(b) By calculating the PDF on a grid of $H$ values, (e.g., <code>np.arange(0, 1.01, 0.01)</code>, calculate the following quantities.

i. posterior mean $\langle H \rangle$ defined as
$$
\langle H \rangle = \int_0^1 H p(H|D,I) dH
$$

ii. posterior median $\tilde H$ such that

$$
\int_0^{\tilde H} p(H|D,I) dH = 0.5
$$

iii. posterior mode where $p(H|D, I)$ is a maximum,

$$
\left.\frac{dp(H|D, I)}{dH} \right|_{H_m} = 0
$$

iv. credible $C$ = 90% region $R$ as defined in the textbook, i.e., provide the lower
and upper bounds $H_l$ and $H_u$ for $R$ such that

$$
\int_{H_l}^{H_u} p(H|D,I) dH = 0.90
$$

<!-- grade_id: cell-dc4933f95a98cb41, points: 2, grade: true, solution: true -->
Question 3a:
    
YOUR ANSWER HERE

<!-- grade_id: cell-a4aa8b2bfc50807c, points: 0, grade: false, solution: true -->
```python
# Question 3b
import scipy

def posterior_quantities(C=0.90):
    """Returns 5 quantities (mean, median, mode, lower, upper) in a tuple
    for the posterior probability distribution:
    
        p(H|D,I) = A H^2 (1-H)

    where mean, median, mode are the corresponding values for this distribution
    and (lower,upper) are the bounds of the 90% (or more generally C) confidence region.
    """
    # Hint: Create a numpy array of H values e.g. H = np.arange(0, 1.01, 0.01)
    #       Compute the correponding p(H) array using numpy math.
    #       Then calculate the values from this array.
    # Hint 2: The function scipy.integrate.cumtrapz may be useful...
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-39ca120f8defa642, points: 5, grade: true, solution: false -->
```python
mean, median, mode, lower, upper = posterior_quantities()

print(f'The mean is {mean}')
print(f'The median is {median}')
print(f'The mode is {mode}')
print(f'The 90% confidence interval is ({lower}, {upper})')

```

```python

```