**Question 1**

(This is Problem 4.4 from Gregory.) A bottle contains three green balls and three red balls.
The bottle is first shaken to mix up the balls.

(a) (2 pts) Calculate analytically the probability that blindfolded, you will pick a red ball
on the third pick, if you learn that at least one red ball was picked on the first two picks.

(b) CODING: (3 pts) Write a Python function that simulates this an instance of drawing
three balls from the bottle.

(c) CODING: (3 pts) Using the function from b), empirically calculate the probability
from part a).

<!-- grade_id: cell-a8606aa986a5e617, points: 2, grade: true, solution: true -->
Quesiton 1a:

YOUR ANSWER HERE

<!-- grade_id: cell-a29bdea1673f34c6, points: 0, grade: false, solution: true -->
```python
# Question 1b

import numpy as np

def simulate_draw3():
    """Simulate drawing 3 balls from a bottle containing 3 green balls and 3 red balls at random
    
    Returns: an array of three strings, each of which is either 'G' or 'R',
             representing green and red respectively.
    """
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-e2d5c2ede65447c8, points: 3, grade: true, solution: false -->
```python
draw = simulate_draw3()

# If you run this multiple times, this should (usually) change each time.
print(f'The three draws are: {draw}')

# Some sanity checks:

# Length should be 3
assert len(draw) == 3

# Each one should be either 'G' or 'R'
d1, d2, d3 = draw
assert d1 in ['G', 'R']
assert d2 in ['G', 'R']
assert d3 in ['G', 'R']

```

<!-- grade_id: cell-fbb32c8fb25debed, points: 0, grade: false, solution: true -->
```python
# Question 1c

def estimate_prob():
    """Estimate the probability of drawing a red ball 3rd given that at least one of the first two
    draws were red, by runing simulate_draw3 many times.
    
    Returns 
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-8832a8df51899375, points: 3, grade: true, solution: false -->
```python
prob = estimate_prob()
print(f'The estimated probability is {prob}')

# You should compare this to what you got in Question 1a.  It should be equal to 2 decimal places.

```

**Question 2**

(6 pts; this is Problem 4.7 from Gregory.) In a particular water sample, ten bacteria are
found, of which three are of type $A$. Calculate analytically the probability of obtaining six
type $A$ bacteria, in a second independent water sample containing 12 bacteria in total.

<!-- grade_id: cell-446be4841a80674a, points: 2, grade: true, solution: true -->
Hint: You may find the following definite integral to be helpful:

$$
\int_0^1 dx x^k (1-x)^{n-k} = \frac{1}{(n+1)\binom{n}{k}}
$$

YOUR ANSWER HERE

**Question 3**
    
This problem demonstrates empirically several properties of counting statistics, which obeys the 
Poisson distribution. We will approximate a Poisson process by generating random values between
0 and 10, but only looking at the ones between 0 and 1.  Technically, this is a multinomial distribution,
but for large $N$ and at least moderately large number of bins, it is a good approximation to
a Poisson process.

(a) (3 pts) Generate $N = 10^6$ uniform random deviates between 0 and 10.
    Count the number of deviates that fall in each of 100 equal-sized bins between 0 and 1 
    (0.00-0.01, 0.01-0.02, etc).  Calculate the mean $\mu$ and variance $\sigma^2$ of the counts per bin.
    Make sure that the mean is almost exactly what you expect: $\bar{n} \approx N/100 = 1000$.

(b) (3 pts) Repeat part (a) for 10 and 1000 equally-sized bins.  This empirically shows that $\mu = \sigma^2 = \bar{n}$, 
    which is an important property of the Poisson distribution.
    
(c) **EXTRA CREDIT** (2 pts): Now generate $M = 100$ realizations of part (a).
    Make sure to not use the same random seed each time.  You should now have $M = 100$ values of both the mean and 
    variance.  Let's call these $\mu_i$ and $\sigma_i^2$, respectively, where $i = 1, 2, ..., M$.
    Now calculate the mean and variance of $\mu_i$ and $\sigma_i^2$.  This demonstrates that the general property 
    that the variance (2nd moment) is much more difficult to measure accurately than the mean (1st moment).

<!-- grade_id: cell-30668e3a915c451c, points: 0, grade: false, solution: true -->
```python
# Question 3a,b

def make_counts(ntot, nbins):
    """Make a random array of ntot uniform random deviates between 0 and 10, 
    and bin the values between 0 and 1 into nbins bins.
    
    Returns the number of values in each bin as an array of lenth nbins
    """
    # Hints:
    # 1. Use either np.random.random or np.random.uniform to make the full array.
    # 2. There are a number of ways to bin these.  Probably np.histogram is easiest.
    #    Read the docs carefully, since you'll need to set some non-default parameters.
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
def calculate_mean_var(x):
    """Return the mean and variance of the values in the given array (as a tuple).
    """
    # You've done this twice already, so this should be a trivial one-liner at this point.
    # Feel free to use numpy functions for this.
    
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-78364ea94b9eec92, points: 6, grade: true, solution: false -->
```python
N = 1_000_000 # The underscores are ignored. For large integers, they are allowed to make it easier to read.

# 4a:
nbins = 100

x = make_counts(N, nbins)
mean, var = calculate_mean_var(x)

print(f"Mean of {nbins} values = {mean:.2f}")
print(f"Expected mean = {N/10/nbins:.2f}")

print(f"Variance of {nbins} values = {var:.2f}")
print(f"Expected variance = {N/10/nbins:.2f}")

# 4b:
nbins = 10
x = make_counts(N, nbins)
mean, var = calculate_mean_var(x)

print()
print(f"Mean of {nbins} values = {mean:.2f}")
print(f"Expected mean = {N/10/nbins:.2f}")

print(f"Variance of {nbins} values = {var:.2f}")
print(f"Expected variance = {N/10/nbins:.2f}")

nbins = 1000
x = make_counts(N, nbins)
mean, var = calculate_mean_var(x)

print()
print(f"Mean of {nbins} values = {mean:.2f}")
print(f"Expected mean = {N/10/nbins:.2f}")

print(f"Variance of {nbins} values = {var:.2f}")
print(f"Expected variance = {N/10/nbins:.2f}")



```

<!-- grade_id: cell-6b7208ff159c9567, points: 0, grade: false, solution: true -->
```python
# Problem 3c:

def run_simulation(M, N, nbins):
    """Run M realizations of the calculation done for Problem 4a
    
    Returns two lists (mu, sigsq), each of length M.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ffbd9ac6df75e7a0, points: 2, grade: true, solution: false -->
```python
mu, sigsq = run_simulation(100, 1_000_000, 100)

mu_mean, mu_var = calculate_mean_var(mu)
sigsq_mean, sigsq_var = calculate_mean_var(sigsq)

print(f"The mean and variance of mu = {mu_mean:.2f}, {mu_var:.2f}")
print(f"The mean and variance of sigsq = {sigsq_mean:.2f}, {sigsq_var:.2f}")

print(f"I.e. the mean values fall mostly in the range {mu_mean:.2f} +- {np.sqrt(mu_var):.2f}")
print(f"and the sigsq values fall mostly in the range {sigsq_mean:.2f} +- {np.sqrt(sigsq_var):.2f}")

```

```python

```