As you will see, we are introducing `pandas` in this assignment.  Pay close attention to the hints.  You can also use `numpy`, `scipy`, and `matplotlib`, but do not use any fitting packages such as `scipy.optimize.curve_fit` to solve for the parameters and uncertainties.  As the fitting functions are linear, the solutions can be obtained entirely with `numpy`'s linear algebra operations.

---

**Question 1**

**CODING:** (3 pts) The file `SunspotNumber.dat` contains data on the number of sunspots observed each day since January 1, 1818 through September 30, 2020.  Columns are 1) year, 2) month, 3) day, and 4) number of sunspots observed (`-1` indicates that no data were taken on that day -- i.e., the data is missing).

```
...
1819 03 21    0
1819 03 22   30
1819 03 23   20
1819 03 24   -1
1819 03 25   -1
1819 03 26   17
...
```

(a) Calculate the monthly mean of the sunspot number. Ignore uncertainties. The output should be a pandas Series indexed by (year, month) with the mean number of sunspots in each month.

(b) Repeat part (a) for the yearly mean.  This will be used in the next problem.

```python
import pandas as pd

sun_df = pd.read_csv('SunspotNumber.dat', 
                     names=['year', 'month', 'day', 'sunspots'],  # Give the names of the columns
                     delim_whitespace=True,  # The default is to use ',' as the delimiter. 
                     na_values=-1,  # Tell pandas that -1 means No Data.
                    )
sun_df.describe()
```

<!-- grade_id: cell-d74f9b41ff3251f1, points: 0, grade: false, solution: true -->
```python
# Question 1a
def mean_sunspots_by_month(df):
    """
    Calculate the mean number of sunspots observed in each month.

    The input pandas data frame include columns named 'year', 'month', 'day', and 'sunspots'.
    
    Rows where the number of sunspots is NaN should be ignored when taking the average.
    
    Returns a pandas Series indexed by (year, month) with the mean number of sunspots in each month.
    """
    # Hints: 
    # 1. Use the groupby method to group the rows by year and month values.
    # 2. Select the 'sunspots' column from this structure.
    # 3. Then the mean method will compute the mean in each group.
    
    # YOUR CODE HERE
    raise NotImplementedError()

# Question 1b
def mean_sunspots_by_year(df):
    """
    Calculate the mean number of sunspots observed in each year.

    The input pandas data frame include columns named 'year', 'month', 'day', and 'sunspots'.
    
    Rows where the number of sunspots is NaN should be ignored when taking the average.
    
    Returns a pandas Series indexed by (year) with the mean number of sunspots in each year.
    """
    # Hint: Same as above, but grouped by year only.
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-26a4e5c807783b2c, points: 3, grade: true, solution: false -->
```python
import numpy as np

sunspots_by_month = mean_sunspots_by_month(sun_df)
print('Sunspots by month:\n')
print(sunspots_by_month,'\n')
print(sunspots_by_month.describe())

sunspots_by_year = mean_sunspots_by_year(sun_df)
print('\n\nSunspots by year:\n')
print(sunspots_by_year,'\n')
print(sunspots_by_year.describe())

# Some sanity checks
assert len(sunspots_by_year) == len(np.unique(sun_df['year']))
assert np.all(sunspots_by_year.index == np.unique(sun_df['year']))
assert sunspots_by_year.index[0] == np.min(sun_df['year'])
assert sunspots_by_year.index[-1] == np.max(sun_df['year'])

assert np.ceil(len(sunspots_by_month) / 12) == len(sunspots_by_year)
assert sunspots_by_month.index[-1] == (2020, 9)

```

---

**Question 2**

The file `SatelliteReentry.dat` contains the number of satellites in low-Earth orbits that reentered and burned in the Earth's atmosphere from 1969 to 2004.

```
1969   26
1970   25
1971   19
1972   12
1973   14
1974   21
...
```

The goal of this problem is to see if there is a relation between solar activity (using sunspot number as a proxy) and the number of satellites that reenter the atmosphere.  We will model the relation with a straight line given by:

\begin{equation}
  N_{\text{reentry}} = a + b N_{\text{sunspot}}
\end{equation}

where $a$ and $b$ are the fitting parameters.


a) (3 pts) Taking the gaussian standard deviation of $N_i$ to be $\sqrt{N_i}$, determine the maximum likelihood estimate of $a$ and $b$.
    
b) (3 pts) Determine their standard deviations $\sigma_a$ and $\sigma_b$, and their covariance $\sigma_{ab}$.
    
c) (1 pt) Plot the data with uncertainties and the best-fit model superimposed.  Make the plot look nice - i.e., label axes, use legible markers, colors, etc.
    
d) (2 pts) Calculate the $\chi^2$ values on a grid of $a$ and $b$ and use the matplotlib `contour` function to plot the contours of constant $\chi^2$ values for $\chi^2 = \chi^2_{\mathrm{min}} + 2.30, 6.17,$ and $11.8$.  Make sure your grid and contour plot is large enough to show all three contour levels.

```python
sat_df = pd.read_csv('SatelliteReentry.dat',
                     names=['year', 'reentries'],  # Give the names of the columns
                     delim_whitespace=True,  # The default is to use ',' as the delimiter. 
                    )
sat_df.describe()
```

<!-- grade_id: cell-ee34684abe2e21be, points: 0, grade: false, solution: true -->
```python
# Question 2a
def fit_reentries_vs_sunspots(n_sunspots, n_reentries):
    """
    Fit a linear regression to the relation
    
        n_reentries = a + b n_sunspots
        
    The inputs are pandas Series instances, indexed by year.  However, they don't (necessarily) have the
    same years, so this function will only consider the subset of years common to both series.
    
    The uncertainty in the number of reentries is taken to be sigma_N = sqrt(N).
    
    Returns a, b
    """
    # Hints: 
    # 1. You can do math with pandas series just like numpy arrays.  
    #    They can also be used as arguments to (most) numpy functions where an array could be used.  
    #    So you can write x**2, x*y, np.sum(x), etc.
    # 2. Make sure to use just the subset of years where you have both kinds of data.
    #    e.g. with the pandas Index.intersection method.
    # 3. For the MLE solution, you need to code this up yourself.  Don't use a canned fitting function from
    #    scipy, scikit-learn, or other similar package.

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3a1eb4e0f7ecae2a, points: 3, grade: true, solution: false -->
```python
n_sunspots = mean_sunspots_by_year(sun_df)
n_reentries = sat_df.set_index('year')['reentries']

a, b = fit_reentries_vs_sunspots(n_sunspots, n_reentries)

print("Best fit solution is Nr = a + b Ns = {:.3f} + {:.3f} Ns".format(a,b))

```

<!-- grade_id: cell-4ce01f76da00b16c, points: 0, grade: false, solution: true -->
```python
# Question 2b
def calculate_cov(n_sunspots, n_reentries):
    """Calculate the covariance matrix of (a,b), the linear fit to
    
        n_reentries = a + b n_sunspots
    
    Returns 2D numpy array [[ sigma_a^2  sigma_ab  ]
                            [ sigma_ab   sigma_b^2 ]]
    """
    # Note: Again, code this up yourself.  Don't use a canned fitting function from
    #       scipy, scikit-learn, or other similar package.
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4470945549496a11, points: 3, grade: true, solution: false -->
```python
cov = calculate_cov(n_sunspots, n_reentries)
print("a = {:.3f} +- {:.3f}".format(a, np.sqrt(cov[0,0])))
print("b = {:.3f} +- {:.3f}".format(b, np.sqrt(cov[1,1])))
print("Covariance of (a,b) = {:.3f}".format(cov[0,1]))

```

<!-- grade_id: cell-675f325c374d120f, points: 1, grade: true, solution: true -->
```python
import matplotlib
import matplotlib.pyplot as plt
%matplotlib inline

# Question 2c
def plot_linear_fit(ax, n_sunspots, n_reentries, a, b, cov):
    """Plot the linear fit n_reentries = a + b n_sunspots on the given matplotlib axis.
    """
    # Hints:
    # 1. Plot all the points (n_sunspots, n_reentries) with error bars for n_reentries
    #    based on sigma(N) = sqrt(N)
    # 2. Draw the best linear fit N = a + b T
    # 3. Write the best fit in text on the plot somewhere
    # 4. Don't forget to label your axes.
    
    # YOUR CODE HERE
    raise NotImplementedError()

fig, ax = plt.subplots(1,1, figsize=(14,8))
plot_linear_fit(ax, n_sunspots, n_reentries, a, b, cov)
plt.show()
```

<!-- grade_id: cell-d607f1317364213f, points: 0, grade: false, solution: true -->
```python
# Question 2d

# First a helper function to calculate the chisq value for any proposed values of a,b.
def calculate_chisq(n_sunspots, n_reentries, a, b):
    """Calculate chisq for a given proposed solution (a, b), assumuing var(N) = N.

    chisq = Sum_i (Nr_i - N_fit(Ns_i))^2 / var(Nr_i),
    where N_fit(Ns) = a + b Ns
    
    Returns chisq for the proposed values of a and b.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e4e397c46983f9e8, points: 1, grade: true, solution: false -->
```python
# Let's check this function at a few values:

# First, the value that we think is the minimum
a, b = fit_reentries_vs_sunspots(n_sunspots, n_reentries)
bestfit_chisq = calculate_chisq(n_sunspots, n_reentries, a, b)
print('chisq at best fit solution = {:.3f}'.format(bestfit_chisq))

# Now a few values a bit shifted from what should be the minimum chisq
chisq1 = calculate_chisq(n_sunspots, n_reentries, a+2, b)
chisq2 = calculate_chisq(n_sunspots, n_reentries, a-2, b)
chisq3 = calculate_chisq(n_sunspots, n_reentries, a, b+0.05)
chisq4 = calculate_chisq(n_sunspots, n_reentries, a, b-0.05)
print('chisq at some nearby values = {:.2f}, {:.2f}, {:.2f}, {:.2f}'.format(chisq1,chisq2,chisq3,chisq4))

# These should all be larger than the best fit chisq
assert chisq1 > bestfit_chisq
assert chisq2 > bestfit_chisq
assert chisq3 > bestfit_chisq
assert chisq4 > bestfit_chisq

# Even a small shift from the best fit should give a slightly larger chisq
chisq5 = calculate_chisq(n_sunspots, n_reentries, a+0.05, b)
chisq6 = calculate_chisq(n_sunspots, n_reentries, a-0.05, b)
chisq7 = calculate_chisq(n_sunspots, n_reentries, a, b+0.001)
chisq8 = calculate_chisq(n_sunspots, n_reentries, a, b-0.001)
print('chisq at some very nearby values = {:.3f}, {:.3f}, {:.3f}, {:.3f}'.format(chisq5,chisq6,chisq7,chisq8))

# These should all still be slightly larger than the best fit chisq
assert chisq5 > bestfit_chisq
assert chisq6 > bestfit_chisq
assert chisq7 > bestfit_chisq
assert chisq8 > bestfit_chisq

```

<!-- grade_id: cell-55fdff15450caffb, points: 1, grade: true, solution: true -->
```python
# Question 2d (continued)

# Now we can use the above function to plot chisq values for many a,b values as a contour plot.
def plot_chisq_ellipses(ax, n_sunspots, n_reentries):
    """Plot ellipses at the 1, 2, and 3-sigma contours of a,b
    """
    # Hints: 
    # 1. Use the ax.contour function.
    # 2. Use fit_reentries_vs_sunspots to determine where the center of the plot should be.
    # 3. Use calculate_cov to get an estimate of the two sigmas, so you can have
    #    the ranges of a,b go to about +- 4 sigma in each direction.
    # 4. If your contours look too chunky, try increasing the number of points.
    
    # YOUR CODE HERE
    raise NotImplementedError()

fig, ax = plt.subplots(1,1, figsize=(14,8))
plot_chisq_ellipses(ax, n_sunspots, n_reentries)
plt.show()
```

---

**Question 3**

It is common practice in data analysis to demonstrate that your estimator produces unbiased results - i.e., the true parameters values are *on average* recovered.  Do the following to show that your estimators for $a$ and $b$ above are unbiased:

(a) (3 pts) Write a function that draws $N$ points $\{x_i, y_i\}$ from the model,
$$
    y = a + b x
$$
where $x_i$ is drawn from a uniform distribution in the range $[x_{\mathrm{min}}, x_{\mathrm{max}}]$ and $y_i$ is drawn from a Gaussian distribution centered on $y = a + bx_i$ with fixed standard deviation $\sigma$.  Also, write a function that fits for $a, b$ given arrays $x, y$.
  
(b) (3 pts) Using $N=90$, $x_{\mathrm{min}}=18, x_{\mathrm{max}}=38, a=200, b=10$, and $\sigma=10$, generate 10000 realizations of the dataset, determine the inverse-variance-weighted mean values of $\langle a \rangle$ and $\langle b \rangle$, along with their uncertainties $\sigma_{\langle a \rangle}$, $\sigma_{\langle b \rangle}$.  Show that your estimators for $a$ and $b$ are unbiased - i.e, the true value falls within one or two standard deviations from the means of $a$ and $b$.

(c) (2 pts) Plot the histogram of $\chi^2$ values of the 10,000 realization from above, and overplot an appropriately scaled $p(\chi^2; \nu)$ distribution for $\nu=90-2=88$ degrees of freedom.

<!-- grade_id: cell-fa66d1ca9e156b3f, points: 0, grade: false, solution: true -->
```python
# Question 3a
def sim_linear(N, xmin, xmax, a, b, sigma):
    """Simulate an experiment with the true relation y = a + b x.
    
    N is the number of data points in the returned arrays.
    xmin, xmax give the range of x values.  The returned x values are uniformly sampled from [xmin, xmax]
    a, b are the true coefficients of the linear relation y = a + b x
    sigma is a constant Gaussian uncertainty of the y values relative to the truth values.

    Returns x, y, both N-element numpy arrays
    """
    # Hint: Use numpy.random.normal for the Gaussian distribution.
    
    # YOUR CODE HERE
    raise NotImplementedError()

def fit_linear(x, y, sigma):
    """Perform a linear fit y = a + b x, given a constant uncertainty, sigma, of the y values.
    
    Returns a, b, sigma_a, sigma_b
    """
    # Note: Don't use a canned fitting function from scipy, scikit-learn, or other similar package.

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-15e11e4acecdced8, points: 3, grade: true, solution: false -->
```python
N = 90
xmin = 18
xmax = 38
true_a = 200
true_b = 10
sigma = 10

# Before running a full simulation suite, it's a good idea to check the two ancillary functions
# to make sure they are working correctly.
# Make a single simulation and fit the parameters for it to make sure you get something reasonable.
x, y = sim_linear(N, xmin, xmax, true_a, true_b, sigma)
print('One simulation yielded arrays:')
print('x =',x)
print('y =',y)

assert len(x) == N
assert len(y) == N
assert np.all(x >= xmin)
assert np.all(x <= xmax)

# And check that your fitting function gets something close to the true values.
a, b, siga, sigb = fit_linear(x, y, sigma)
print('The linear fit for this simulation is:')
print('a = {:.3f} +- {:.3f}'.format(a, siga))
print('b = {:.3f} +- {:.3f}'.format(b, sigb))

# These aren't necessarily very close to the true values, but they should be
# vaguely similar.  (We'll check this more precisely in 3b)
assert np.isclose(a, true_a, rtol=0.1)
assert np.isclose(b, true_b, rtol=0.1)

# These aren't techinically *always* true, but the estimated values should almost always be within
# 5 sigma of the true answer.  Indeed >99% of the time, they should be within 3 sigma.
assert abs(a - true_a) < 5*siga
assert abs(b - true_b) < 5*sigb

```

<!-- grade_id: cell-c6e9f5a3e893f49a, points: 0, grade: false, solution: true -->
```python
# Question 3b
def run_simulation(N, xmin, xmax, true_a, true_b, sigma, Nsim):
    '''
    1. Generates Nsim realizations of the dataset.
    2. Determine the MLE of a, b, sigma_a, sigma_b for 
       each realization
    3. Find the chisq of each realization
    4. Calculate the inverse-variance-weighted mean of
       a and b and the corresponding uncertainties
       for each realiation.
    
    Returns: chisqs, mean_a, mean_b, sig_meana, sig_meanb
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-088e6015d3f58554, points: 3, grade: true, solution: false -->
```python
N = 90
xmin = 18
xmax = 38
true_a = 200
true_b = 10
sigma = 10
Nsim = 10000

chisqs, mean_a, mean_b, sig_meana, sig_meanb = run_simulation(N, xmin, xmax, true_a, true_b, sigma, Nsim)

print('True a, b = ', true_a, true_b)

print('Inverse-variance-weighted mean of sims:')
print('   a = {:.5f} +- {:.5f}'.format(mean_a, sig_meana))
print('   b = {:.5f} +- {:.5f}'.format(mean_b, sig_meanb))

assert len(chisqs) == Nsim

print('Mean chisq of simulations is {:.2f}'.format(np.mean(chisqs)))

# Now our estimates should be much more precise.
assert np.isclose(mean_a, true_a, rtol=1.e-3)
assert np.isclose(mean_b, true_b, rtol=1.e-3)

assert abs(mean_a - true_a) < 5*sig_meana
assert abs(mean_b - true_b) < 5*sig_meanb

```

<!-- grade_id: cell-d6a16b62ce735228, points: 2, grade: true, solution: true -->
```python
import scipy.stats

# Question 3c
def plot_chisq_hist(ax, chisq, nu):
    '''Plots a histogram of chisq values
    Also overplots a chisq distribution for nu degrees of freedom.
    '''
    # Hint: See the pdf method of scipy.stats.chi2.
    
    # YOUR CODE HERE
    raise NotImplementedError()

fig, ax = plt.subplots(1, 1, figsize=(14,8))
plot_chisq_hist(ax, chisqs, N-2)
plt.show()
```

```python

```