**Problem 1**

**CODING:** This is a continuation of HW7 where you will learn how to perform non-linear regression and also see how more data improves the statistical quality of the data and the resulting parameter constraints.  There are 10 data files `run[0-9].dat` that contain hypothetical particle-physics spectral data (energy vs counts) of a search for a new particle called the *Riggs boson*.  The data were taken independently in 10 distinct runs.  The columns are 1) energy $E$ (GeV) and 2) counts $N$ (number of events in that energy bin).
  
       1.0    96
       2.0    99
       3.0   111
       4.0   124
       5.0   105
       6.0   102
       ...

  
  A signal of a new particle consists of a gaussian on top of a smooth background.  Model the spectrum with a function given by:
$$
  N(E) = a + bE +cE^2 + Ae^{-(E-E_{\mathrm{Riggs}})^2/(2\sigma_E^2)}
$$
where we now have 6 fit parameters -- $a,b,c,A,E_{\mathrm{Riggs}},\sigma_E$.  The first three terms make up the background, the last term represents the signal.  We assume here that $\sigma_E$ and $E_{\mathrm{Riggs}}$ are unknown, but they are expected to be in the range 1 - 5 GeV and 40 - 90 GeV, respectively.

a) (4 pts) In this first part, use only the first data file `run0.dat`.  Use `scipy.optimize.curve_fit` to determine the best-fit values of the parameters and the covariance matrix.  Can you claim a detection of the Riggs boson?  Justify your answer.

b) (3 pts) Since the 10 datasets are independent, the counts can simply be added to produce a dataset with higher statistical quality.  Successively add the counts in `run1.dat`, `run2.dat`, and so on and repeat the fits.  After which run can you claim a $5\sigma$ detection of the Riggs?

c) (3 pts) Using data from all 10 runs, measure the energy $E_{\mathrm{Riggs}}$ and uncertainty of the the Riggs boson.  What is the final significance (in units of $\sigma$) of the detection?


<!-- grade_id: cell-e9aa02daaf35d396, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd
import scipy.optimize

# Problem 1a
def read_riggs_data(n):
    """
    Read the data file for run n, where 0 <= n <= 9.
    
    The file name is run[n].dat.  (I.e. run0.dat, run1.dat, ... run9.dat)
    
    Returns E, N for this run as numpy arrays
    """
    # Hint: If you use pandas to read the file, to_numpy() will convert to a numpy array

    # YOUR CODE HERE
    raise NotImplementedError()

def riggs_model(E, a, b, c, A, E_Riggs, sigma_E):
    """
    Return the expected counts given the model for the Riggs boson.
    
        N(E) = a + b E + c E**2 + A exp(-(E-E_Riggs)**2/(2 sigma_E**2))
    
    On input, E is an array.  The other parameters are scalars.
    
    Returns N(E) given the model parameters.
    """
    # YOUR CODE HERE
    raise NotImplementedError()

def fit_riggs_model(E, N):
    """
    Find the best fit model given the observed data N(E)

    Returns the parameters [a, b, c, A, E_Riggs, sigma_E] as a numpy array 
    and the covariance matrix, also as a numpy array.
    """
    # Hints:
    # 1. Use scipy.optimize.curve_fit
    # 2. Make sure to set the uncertainties correctly based on Poisson errors (sigma_N = sqrt(N))
    # 3. Set the bounds appropriately for E_Riggs and sigma_E based on the problem statement.
    # 4. For the amplitude, we know it's positive, so [0,1.e10] is appropriate.
    # 5. For the first three [-1.e10,1.e10] is fine. (i.e. basically no prior information, although
    #    it is helpful to put some kind of upper bounds other than infinity).
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
def calculate_significance(params, cov):
    """
    Calculate the significance of a possible detection of the Riggs boson.
    
    params = [a, b, c, A, E_Riggs, sigma_E]
    cov = the estimated covariance matrix
    
    Returns the S/N of the purported detection.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
def calculate_energy(params, cov):
    """
    Calculate the estimated value of E_Riggs and its uncertainty.

    params = [a, b, c, A, E_Riggs, sigma_E]
    cov = the estimated covariance matrix
    
    Returns the estimate of (E_Riggs, sigma(E_Riggs)).
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0dcecef09a7cd732, points: 3, grade: true, solution: false -->
```python
import matplotlib
import matplotlib.pyplot as plt
%matplotlib inline

E, N = read_riggs_data(0)

params, cov = fit_riggs_model(E,N)

print('params = ',params)
print('sigma = ',np.sqrt(np.diag(cov)))
print('Significance = ',calculate_significance(params, cov))

# Quick and dirty visualization of the data and the fit
fig, ax = plt.subplots(1,1, figsize=(14,8))
ax.plot(E,N)
ax.plot(E,riggs_model(E,*params))
ax.set_xlabel('E')
ax.set_ylabel('N')
plt.show()
```

<!-- grade_id: cell-4bce5209d816642f, points: 1, grade: true, solution: true -->
(Problem 1a) 

Can you claim a detection of the Riggs boson?  Justify your answer.

----

YOUR ANSWER HERE


```python
fig, ax = plt.subplots(1,1, figsize=(14,8))

E, N = read_riggs_data(0)
params, cov = fit_riggs_model(E,N)
nu = calculate_significance(params, cov)
E_R, sig = calculate_energy(params, cov)
print('i=0: nu = {}, E = {} +- {}'.format(nu,E_R,sig))

for i in range(1,10):
    #print('N  = ',N)
    Ei, Ni = read_riggs_data(i)
    #print('Ni = ',Ni)
    assert np.allclose(E,Ei)
    N += Ni

    params, cov = fit_riggs_model(E,N)
    nu = calculate_significance(params, cov)
    E_R, sig = calculate_energy(params, cov)
    print('i={}: nu = {}, E = {} +- {}'.format(i,nu,E_R,sig))

    ax.plot(E,N)
    ax.plot(E,riggs_model(E,*params))

ax.set_xlabel('E')
ax.set_ylabel('N')
plt.show()
```

<!-- grade_id: cell-a29e6144b4c98700, points: 3, grade: true, solution: true -->
(Problem 1b) 

After which run can you claim a 5 sigma detection of the Riggs?

YOUR ANSWER HERE

<!-- grade_id: cell-6f76f7f3303d1634, points: 3, grade: true, solution: true -->
(Problem 1c) 

Using data from all 10 runs, measure the energy E_Riggs and uncertainty of the the Riggs boson. 

YOUR ANSWER HERE

----

What is the final significance of the detection?

YOUR ANSWER HERE

**Problem 2**

**CODING:** Markov Chain Monte Carlo (MCMC) is a powerful tool for sampling the posterior distribution of complicated models.  The best way to learn how to implement this is through an example,.  In this problem, you will redo parts of HW6 Problem 2) using the Metropolis-Hastings MCMC algorithm.

Recall that we were tasked to model the relation between the number of satellites that reenter the Earth $N_{\mathrm{reentry}}$ and the average number of sunspots $N_{\mathrm{sunspot}}$ in a given year.  The data are provided in `ReentryData.dat`.  We used the following model:
$$
  N_{\text{reentry}} = a + b N_{\text{sunspot}}
$$
where $a$ and $b$ are the fitting parameters.  Using the analytic solution to the MLE, we saw that the posterior distribution is characterized by the means and covariance matrix given by:
$$
\begin{align}
    \left( \begin{array}[r] ~ a \\ b \end{array} \right)
    &=
    \left( \begin{array}[r] ~ 13.11 \\ 0.110 \end{array} \right) &
    \Sigma &=
    \left( \begin{array}[rr] ~ 1.84 & -0.0141 \\ -0.0141 & 0.000169 \end{array} \right)
\end{align}
$$
  

a) (8 pts) Construct a Markov Chain as follows:
   * Initialization:  Start with the best-fit parameters $\vec{x}_{\mathrm{old}} = (13.11~0.110)$.  Calculate the value of the likelihood $\mathscr{L}(\vec{x}_{\mathrm{old}})$ at this point.  The step sizes $\delta\vec{x}$ must be chosen appropriately such that 
    
   * For each iteration:
     * Generate a candidate random step from $\vec{x}_{\mathrm{old}}$ by drawing a pair of random numbers $\vec{r}$ from a normal distribution $\mathscr{N}(0, \delta\vec{x})$.  Your candidate point is $\vec{x}_{\mathrm{new}} = \vec{x}_{\mathrm{old}} + \vec{r}$.
     * Calculate the ratio $R = \mathscr{L}(\vec{x}_{\mathrm{new}})/\mathscr{L}(\vec{x}_{\mathrm{old}})$.
     * Determine whether to accept or reject the candidate step:
       - If $R \geq 1$, take the proposal step (accept).
       - else if $R < 1$, then draw a random number $U$ from a uniform distribution between 0 and 1.
         * If $U < R$, take the proposal step (accept).
         * else if $U \geq R$, do not take the step (reject).
         
   Do $10^5$ iterations and choose $\delta\vec{x}$ such that $\sim 50\%$ ($30-70\%$) of the candidate steps are accepted.  The output should be a Markov chain (list) of accepted steps $\{\vec{x}\}$.
    
b) (4 pts) Using the Markov chain from above, calculate the sample means $(\bar{a}~\bar{b})$ and covariance matrix, and check that the values are nearly identical to those from the analytic MLE.
    
c) (3 pts) Install the plotting package \texttt{corner} in your Anaconda installation.
    
       conda install corner
    
   If you run into problems, see:

   http://corner.readthedocs.io/en/latest/install.html
    
   This package allows you to easily visualize the variance and covariance of an MCMC chain.  A quick-start guide can be found here:

   http://corner.readthedocs.io/en/latest/pages/quickstart.html

   Use this package to visualize the Markov chain, which is just a sampling of the bivariate gaussian characterized by the means and covariance matrix.


<!-- grade_id: cell-b20a447acba1e8b7, points: 0, grade: false, solution: true -->
```python
# Problem 2a

def read_sat_data():
    """Read the satellite and sunspot data and convert sunspots to mean number per year.
    
    Returns n_sunspots, n_satellites as numpy arrays.
    """
    # Hints:
    # 1. Remember to limit the range to the years with data for both sunspots and satellites.
    # 2. Feel free to copy the code from the HW6 solution set.
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
def sat_likelihood(data, theta):
    """Compute the likelihood of our data given a set of parameters theta.
    
    In our case, data is a tuple of numpy arrays: (n_sunspots, n_reentries) and theta is the tuple (a,b).

    The likelihood is a Gaussian, assuming the uncertainty in n_reentries is sqrt(n_reentries).
    And note that since only relative likelihoods for different choices of theta are relevant, 
    you can ignore any constant factors that are independent of theta.
    
    Returns the likelihood for this choice of theta.
    """
    # Hint: This function will be called A LOT during the chain.
    #       You really want to avoid having any loops in this function.
    #       Use numpy array math and things like np.sum to avoid explicit loops.
    #       If you don't know how to do this, you might want to look at how similar 
    #       functions were implemented in past solution sets.
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2499fb2af20e4364, points: 3, grade: true, solution: false -->
```python
import time

n_sunspots, n_reentries = read_sat_data()

print('n_sunspots = ',n_sunspots)
print('n_reentries = ',n_reentries)
print('len = ',len(n_sunspots),len(n_reentries))

# Make sure the data are clipped to the same time period.
assert len(n_sunspots) == len(n_reentries)

# Make sure we convert to numpy arrays if necessary.
assert type(n_sunspots) == np.ndarray
assert type(n_reentries) == np.ndarray

# Check our likelihood function
data = (n_sunspots, n_reentries)
mle_theta = (13.11, 0.110)  # This is approximately the MLE solution from HW6.
print('Likelihood at MLE solution = ',sat_likelihood(data, mle_theta))

# This is expected to be the maximum likelihood.  Make sure it is a peak.
print('Likelihood at some other nearby locations:')
for theta in [ (13.05, 0.11), (13.15, 0.11), (13.11, 0.105), (13.11, 0.115)]:
    print('   L({}) = {}'.format(theta, sat_likelihood(data, theta)))
    assert sat_likelihood(data, theta) < sat_likelihood(data, mle_theta)

# Finally, make sure our likelihood function is fast enough to run many thousands of times in the MCMC.
t0 = time.time()
for i in range(100):
    theta = np.random.random(2) * 10
    sat_likelihood(data,theta)  # ignore return value
t1 = time.time()
print('time to run sat_likelihood 100 times = ',t1-t0)

t2 = time.time()
for i in range(10000):
    theta = np.random.random(2) * 100
    sat_likelihood(data,theta)  # ignore return value
t3 = time.time()
print('time to run sat_likelihood 10,000 times = ',t3-t2)
assert t3-t2 < 5
# If this is more than 5 seconds, you need to work on speeding it up.
# Note: for the official solution, this takes under 0.1 second on a modern laptop.
```

**Aside:** *Quick overview of Python Classes*

Those of you who have used Python classes before, feel free to skip this section.  For the rest of you, this should hopefully give you the basics of what you need to know about classes to complete this assignment.

A class (in any object-oriented language, not just Python) is a data structure that holds some combination of data, typically stored as attributes, along with some functions that can act on those data.

One class you are probably already somewhat familiar with in Python (although you may not have thought of it as a class) is a list.  Every list you make (e.g. `l1 = [1,2,3]`) is an instance of the native Python class `list`.  You can access the class name of any object in Python using a special "dunder" (short for double underscore) attribute `__class__`.  Similarly, you can any object's doc string using `__doc__`:

    >>> l1 = [1,2,3]
    >>> print(l1.__class__)
    <class 'list'>
    >>> print(l1.__doc__)
    Built-in mutable sequence.

    If no argument is given, the constructor creates a new empty list.
    The argument must be an iterable if specified.

The more interesting attributes are functions, which can act on the object's data.  These are called methods in the standard terminology, to distinguish from free functions, not embedded in a class.  Some methods of `list` that you might be familiar with include `sort`, `append`, `reverse`, and `clear`:

    >>> l2 = [5, 1, 12, 8, 0]
    >>> l2.sort()
    >>> print(l2)
    [0, 1, 5, 8, 12]
    >>> l2.append(6)
    >>> print(l2)
    [0, 1, 5, 8, 12, 6]
    >>> l2.reverse()
    >>> print(l2)
    [6, 12, 8, 5, 1, 0]
    >>> l2.clear()
    >>> print(l2)
    []

Python makes it easy to define your own classes, which we will do for this problem.  We'll define a class called `MCMC`, which we'll use to keep track of the samples of the Markhov Chain as we run through a bunch of steps.  We will make an instance of this class according to a number of parameters:

    sat_mcmc = MCMC(sat_likelihood, data, theta, step_size, names)
    
This looks a bit like a function call, but the right hand side is the name of the class, followed by the parameters we'll use when making the class.  The initialization is actually done by a method called `__init__`, which is defined in the class definition.  More details on what all these parameters mean are given in the doc string below.

The return value, which we named `sat_mcmc` is an instance of the `MCMC` class, which we can use to run the chain by calling various methods:

    sat_mcmc.burn(100)
    sat_mcmc.run(10000)
    
and then make some visualizations to see how the run is progressing:

    sat_mcmc.plot_samples()
    sat_mcmc.plot_hist()

These methods are all already defined for you, since they are fairly straightforward.  The real meat of the calculation is in the `step` method, which carries out one step of the chain.  This method you will need to write yourself.  

You'll note that the first argument of each method defined below is a special parameter called `self`.  This is the name of the instance that was used to call the method.  E.g. when you call `sat_mcmc.burn(100)`, this turns into a call to `MCMC.burn` with the first parameter being `self=sat_mcmc` and the second parameter being `nburn=100`.  This lets you access the stored attributes of current instance from inside the function definition (e.g. those that are set in the `__init__` like `self.nparams` or `self.step_size`) or call other methods (e.g. `self.step()`).

When writing the definition of `MCMC.step`, you'll need to use some of the parameters that are saved for you in the initialization, and also update some state parameters to effect a single Metropolis Hastings step.  You can write to any saved attribute just like how you write to a normal variable.  For instance,

    self.naccept += 1
    
will add 1 to the `naccept` attribute of the current instance, which in our case would be `sat_mcmc`.

For another overview of Python classes, here is a pretty good one:

https://www.dataquest.io/blog/using-classes-in-python/

<!-- grade_id: cell-e6b0638b3785a51d, points: 0, grade: false, solution: true -->
```python
# Problem 2a (continued)
class MCMC:
    """Class that can run an MCMC chain using the Metropolis Hastings algorithm

    This class is based heavily on the "Trivial Metropolis Hastings" algorithm discussed in lecture.
    If you haven't used classes before, you can think of it as just a way of organizing the variables
    and functions related to the MCMC operation.

    You use it by creating an instance of the class as follows:

        mcmc = MCMC(likelihood, data, theta, step_size)

    The parameters here are:

        likelihood is a function returning the likelihood p(data|theta), which needs to be
            defined outside the class.  The function should take two variables (data, theta) and 
            return a single value p(data | theta).

        data is the input data in whatever form the likelihood function is expecting it.  
            This is fixed over the course of running an MCMC chain.

        theta is a list or array with the starting parameter values for the chain.

        step_size is a list or array with the step size in each dimension of theta.


    Then once you have an MCMC object, you can use it by running the following functions:

        mcmc.burn(nburn) runs the chain for nburn steps, but doesn't save the values.

        mcmc.run(nsteps) runs the chain for nsteps steps, saving the results.

        mcmc.accept_fraction() returns what fraction of the candidate steps were taken.

        mcmc.get_samples() returns the sampled theta values as a 2d numpy array.


    There are also simple two plotting functions that you can use to look at the behavior of the chain.

        mcmc.plot_hist() plots a histogram of the sample values for each paramter.  As the chain
            runs for more steps, this should get smoother.
        
        mcmc.plot_samples() plots the sample values over the course of the chain.  If the burn in is
            too short, it should be evident as a feature at the start of these plots.


    Finally, there is only one method you need to write yourself.
    
        mcmc.step() takes a single step of the chain.
    """
    def __init__(self, likelihood, data, theta, step_size, names=None, seed=314159):
        self.likelihood = likelihood
        self.data = data
        self.theta = np.array(theta)
        self.nparams = len(theta)
        self.step_size = np.array(step_size)
        self.rng = np.random.RandomState(seed)
        self.naccept = 0
        self.current_like = likelihood(self.data, self.theta)
        self.samples = []
        if names is None:
            names = ["Paramter {:d}".format(k+1) for k in range(self.nparams)]
        self.names = names            

    def step(self, save=True):
        """Take a single step in the chain"""
        # 1. Calculate the new theta value.
        # 2. Calculate the likelihood for that theta.       
        # 3. Decide whether or not to take the step.
        # 4. If taking the step, update self.current_like and self.theta.
        # 5. If save==True, add the sample to self.samples and maybe add 1 to self.naccept.
        
        # Hint: For the random numbers use the stored RandomState, not np.random functions.
        #       Otherwise the comments below about how your results should look might not make much sense.
        #       E.g. self.rng.normal, not np.random.normal.
        #            and self.rng.uniform, not np.random.uniform, etc.
        #       In general, using a well-defined random seed will help keep things repeatable
        #       while still have pseudo-random behavior in your programs.

        # YOUR CODE HERE
        raise NotImplementedError()

    def burn(self, nburn):
        """Take nburn steps, but don't save the results"""
        for i in range(nburn):
            self.step(save=False)

    def run(self, nsteps):
        """Take nsteps steps"""
        for i in range(nsteps):
            self.step()

    def accept_fraction(self):
        """Returns the fraction of candidate steps that were accpeted so far."""
        if len(self.samples) > 0:
            return float(self.naccept) / len(self.samples)
        else:
            return 0.
        
    def clear(self, step_size=None, theta=None):
        """Clear the list of stored samples from any runs so far.
        
        You can also change the step_size to a new value at this time by giving a step_size as an
        optional parameter value.
        
        In addition, you can reset theta to a new starting value if theta is not None.
        """
        if step_size is not None:
            assert len(step_size) == self.nparams
            self.step_size = np.array(step_size)
        if theta is not None:
            assert len(theta) == self.nparams
            self.theta = np.array(theta)
            self.current_like = self.likelihood(self.data, self.theta)
        self.samples = []
        self.naccept = 0
        
    def get_samples(self):
        """Return the sampled theta values at each step in the chain as a 2d numpy array."""
        return np.array(self.samples)
        
    def plot_hist(self):
        """Plot a histogram of the sample values for each parameter in the theta vector."""
        all_samples = self.get_samples()
        for k in range(self.nparams):
            theta_k = all_samples[:,k]
            plt.hist(theta_k, bins=100)
            plt.xlabel(self.names[k])
            plt.ylabel("N Samples")
            plt.show()
        
    def plot_samples(self):
        """Plot the sample values over the course of the chain so far."""
        all_samples = self.get_samples()
        for k in range(self.nparams):
            theta_k = all_samples[:,k]
            plt.plot(range(len(theta_k)), theta_k)
            plt.xlabel("Step in chain")
            plt.ylabel(self.names[k])
            plt.show()
```

<!-- grade_id: cell-3ba647005d2db8eb, points: 3, grade: true, solution: false -->
```python
# It's helpful to start an MCMC at a good initial guess.  The worse the initial guess, the more
# "burn in" is required to produce a good chain.  Starting at the MLE solution means that we shouldn't
# need much burn in, but we'll do a little bit anyway.
theta = mle_theta

# Here, we'll run this with a poor choice of step size to show what that looks like.
# In the next cell down, you can play around with modifying this to produce better results.
step_size = (0.1, 0.1)

# Make the mcmc object with the satellite likelihood function
sat_mcmc = MCMC(sat_likelihood, data, theta, step_size, names=('$a$','$b$'))

# Run the burn-in.  
# This is more important for complicated models where you may not really know a good starting point.
# In that case, a rule of thumb is to burn in for about 10% of the total length of the MCMC run.
# Here, we don't need to be so conservative. 
# Using 100 burn in steps is plenty to get the chain to forget precisely where it started.
sat_mcmc.burn(100)

# In the next section, you will pick a good choice for the number of steps.  Here we do 10000.
sat_mcmc.run(10000)

# Plot the samples to check the burn in and sample size.
sat_mcmc.plot_samples()

# Plot histograms of the parameter values to see if they look converged.
sat_mcmc.plot_hist()

print('After running for 1000 steps:')
print('Acceptance rate is ', sat_mcmc.accept_fraction())

```

Some things to notice about the above plots.

1. The histogram of b looks reasonably Gaussian (albeit noisy).  This is good.
   It means the chain is well sampling the posterior distribution of b around the maximum likelihood value.
2. The samples for b versus step number looks extremely noisy.  This is also good.
   You want the parameter values from one step to the next to jump around a lot relative to the full range of the
   distribution.
3. The histogram of a is not very Gaussian.  It has two peaks and is clearly missing an important part
   of the distribution near the MLE solution.  This means that the chain has not "converged".  It may mean that you
   need to let it run longer, or you may need to adjust the step size.  (Or both.)
4. The samples for a are very correlated from one step to the next.  The a values only slowly meander through the
   allowed space.  The chain happens to have spent most of the time in two portions of the full range, explaining
   why we saw two peaks in the final a histogram.  This is typically a sign that the step size for a is too small.
5. The acceptance rate is only about 0.1.  This is very low.  It means that 9/10 of the candidate steps are rejected.
   This is typically a sign that one or both of the step sizes is too large.  The chain keeps trying to jump too far
   away from the peak, and the step is rejected.
   
Now you have a chance to play around with the two step size parameters to try to pick something more reasonable.
Your goals are:

1. Get the histograms to look reasonably Gaussian (at least single peaked).
2. Get the samples vs step number to look very noisy (not meandering for either one)
3. Get the acceptance rate near 0.5

Once you've done that, feel free to increase the number of samples in the chain.  Don't have it run for more than about 1 minute, but generally the more samples you have, the better the final statistics and contours.

<!-- grade_id: cell-e0f6441d951bd69d, points: 0, grade: false, solution: true -->
```python
# Problem 2a (continued)
def sat_step_size():
    """Return your choice of step size in (a,b) as a tuple."""
    # YOUR CODE HERE
    raise NotImplementedError()

def sat_nsteps():
    """Return how many steps to run"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-907f069ece5e962c, points: 2, grade: true, solution: false -->
```python
theta = mle_theta

sat_mcmc = MCMC(sat_likelihood, data, mle_theta, sat_step_size(), names=('$a$','$b$'))

sat_mcmc.burn(100)

t0 = time.time()
sat_mcmc.run(sat_nsteps())
t1 = time.time()

sat_mcmc.plot_samples()

sat_mcmc.plot_hist()

print('After running for {} steps:'.format(sat_nsteps()))
print('Acceptance rate is ', sat_mcmc.accept_fraction())
print('Time to run chain is {} seconds'.format(t1-t0))

assert 0.3 < sat_mcmc.accept_fraction() < 0.7    # Adjust step sizes if this fails.
assert t1-t0 < 120   # Make sure it is < 60 on your system.  
                     # If it takes longer than 2 min on the TA's computer, we may dock points.
                     # (The official solution gets quite good statistics in only a couple seconds.)

```

<!-- grade_id: cell-08ab7807e5ebc895, points: 0, grade: false, solution: true -->
```python
# Problem 2b
def calculate_mean(mcmc):
    """Calculate the mean of each parameter according to the samples in the MCMC object.

    Returns the mean values as a numpy array.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
def calculate_cov(mcmc):
    """Calculate the covariance matrix of the parameters according to the samples in the MCMC object.

    Returns the covariance matrix as a 2d numpy array.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-91558298b26bef29, points: 4, grade: true, solution: false -->
```python
# Now that we think we have a good chain and that it is converged, we can compute the mean and covariance
# of the parameters.
mean = calculate_mean(sat_mcmc)
cov = calculate_cov(sat_mcmc)

print('Mean values of (a,b) = ', mean)
print('Uncertainties of (a,b) = ', np.sqrt(cov.diagonal()))
print('Covariance matrix = \n', cov)

```

<!-- grade_id: cell-2498d25def3114a9, points: 3, grade: true, solution: true -->
```python
import corner
# Note: You will need to install this.  See http://corner.readthedocs.io/en/latest/install.html

# Problem 2c
def plot_corner(mcmc):
    """Make a corner plot for the parameters a and b with contours corresponding to the same
    delta chisq contours we drew in homework 4.
    """
    # Hint: Use the corner.corner function
    
    # YOUR CODE HERE
    raise NotImplementedError()

plot_corner(sat_mcmc)
```

**Problem 3**

**CODING, EXTRA CREDIT** (3 pts) Repeat Problem 2 for the non-linear regression problem in Problem 1.  
Specifically, generate a 6-parameter Markov Chain, solve for the means and covariance matrix, 
and plot the posterior distributions with $\texttt{corner}$.


<!-- grade_id: cell-781817231b2490a9, points: 3, grade: true, solution: true -->
```python
# Problem 3  (EXTRA CREDIT)

# Hint: With more parameters, the optimal acceptance fraction goes down. 
#       More or less 1/Nparams is good.  So about 0.16 in this case.

# YOUR CODE HERE
raise NotImplementedError()
```

```python

```